# Assignment 3
This assignment consists of two notebook files. Each exercise is graded using *hidden* tests. If you pass these tests then you are rewarded *full points* for that exercise; if your code fails the tests in any way then you will get *no points* for that exercise. Make sure to **read the rules** before you start the assignment.

## Rules
For this assignment the following rules apply:

**General**
 * The assignment should be completed in **groups of two or three** (enroll in a group on Brightspace).
 * Any kind of intergroup discussion will be considered fraud and both the parties will be punished.
 * All code must be written intra group. All external help, with the exception of Python/library documentation and the lecture slides, will be considered fraud (including generative AI).
 * Do not use libraries that implement the assignment for you. Ask a TA if you are unsure.

**Grading**
 * Each exercise is graded using additional hidden tests. These tests will check your solution for different (unseen) inputs.
 * If the tests pass without error (warnings are allowed) then you receive full points.
 * If a test cell takes more than five minutes to complete then this is considered an error.
 * Do not make any assumptions on the input data (such as resolution) unless specified otherwise. Doing so may result in the tests failing and thus 0 points.
 * Your grade is computed as $\frac{\text{points}}{\text{max\_points}} * 9 + 1$ and will be rounded to the closest 0.1 point.
 * Submit your code to Brightspace as a zip file containing only the notebook (`*.ipynb`) files.
 * **Do not rename the notebook files**
 
**Late Submissions**
 * Late submissions must be submitted *as soon as possible* to the "Assignment 3 - Late Submissions" assignment on Brightspace.
 * The following penalty will be applied: $\text{adjusted grade} = \text{grade} - 1 - \lceil\frac{\text{minutes late}}{10}\rceil$

<br />
 
**Before you submit**, make sure that you are not accidentaly using any global variables. Restart the kernel (wiping all global variables) and run the code from top to bottom by clicking "Kernel" => "Restart & Run all" in the menu bar at the top.

```python
%matplotlib widget
import itertools
import matplotlib.pyplot as plt
import numpy as np
import cv2
import scipy
import scipy.linalg
import scipy.signal
import os
# helpers.py is one level up in the directory structure so we need to tell Python were to find it
import sys
sys.path.append("../../")
import helpers
local_data_folder = os.path.join(helpers.dataset_folder, "week3", "collections")

%load_ext autoreload
%autoreload 2
```

## Exercise 0 (0 points)
This is a group assignment. Every student is expected to contribute to all exercises.

Enter the student IDs of all contributing students:

<!-- grade_id: cell-49875c747627d70d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
student_ids = ["1234", "5678"]
```

<!-- grade_id: cell-fc15b7a979cdfe33, points: 0, grade: true, solution: false -->
```python
# DO NOT REMOVE, MODIFY, OR COPY THIS CELL
```

# Searching image collections

In the previous excercises we have seen how an image can be inpainted without access to additional information. Yet, such an approach is limited in the variety of content it can produce. We can achieve much richer and more constent inpainting by copying existing image content from a similar image. For that we need a source of such images - an image collection. Furthermore, we need a fast and reliable way how to find good matches to our query image.

While we could potentially compare images pixel by pixel using the Sum of Squared Distances metric, a more efficient and often more robust approach is to utilize lower-dimensional image descriptors. One of the simplest choices is *GIST*.

The name *GIST* refers to the idea of extracting the gist of the image. GIST is a simple fixed-length vector of real numbers that summarizes distribution of contrast (or gradients in the image). It is built by averaging intensity of edges detected in several image regions, at several scales and several orientations.

In this assignment, we will compute GIST of an image and then use it to find similar images in a preprocessed image collection.

### GIST algorithm
There are several ways how to build a GIST descriptor. In this Assignment we apply a sequence of image convolutions (see Assignment 1.1) with differently oriented and sized Gabor filters. The results of each convolution are then split into a fixed-sized grid and averaged to obtain entries for the descriptor. Finally, a low resolution version of the image is appended to the final descriptor to encode the average color distribution.

## Gabor Filter

The coefficients of the GIST descriptor are composed of image gradients detected at different scales (representing range of spatial frequencies) and different orientations. 

We have previously used specialized filters to detect sharp edges (high frequency) with either horizontal or vertical orientation. 
Gabor filter (or Gabor patches) are a generalization of edge detection that is inspired by processing in human visual cortex.


<img src="gabor.png" alt="Gabor filter" style="width:20%">


A Gabor filter $f(x,y)$ is a product of a symmetrical bi-variate Gaussian filter
$$
g(x,y) = e^{-\frac{1}{2}\frac{x^2+y^2}{\sigma^2}}
$$
and a sinusoidal corrugation
$$
s(x,y) = \sin\left((x \cos(\alpha) - y \sin(\alpha)) \cdot f\right)
$$
with parameters
- $\sigma$ as a standard deviation controlling the width of the filter,
- $\alpha$ as an angle (in radians) controlling the orientation of the filter (and detected gradients)
- $f$ as a spatial frequency controling the sensitivity of the filter to sharp or smooth gradients.

Finally, the entire Gabor filter is
$$
f(x,y) = g(x,y) \cdot s(x,y) \cdot \frac{1}{\sum_{x,y} g(x,y)}
$$
where the denominator on the right normalizes the sum of the Gaussian component to one. 

This normalization ensures that the response of the filter stays normalized regardless of the discretization to fixed-resolution pixel grid of size $s$.

### Tips
Try how the filter changes for different choices of $\sigma$, $\alpha$ and $f$.

```python
def build_gabor_filter(size, sigma, f, alpha):
    x, y = np.meshgrid(np.linspace(-1,1,size),np.linspace(-1,1,size))
    ox = x * np.cos(alpha) - y * np.sin(alpha)
    r2 = x**2 + y**2
    gauss = np.exp(-0.5*r2/sigma**2)
    gauss /= gauss.sum()
    wave = np.sin(ox*f)
    gabor_filter = gauss * wave
    return gabor_filter
    

gabor_filter = build_gabor_filter(35, 1.0/3.2, 8, np.radians(60))
helpers.show_image(gabor_filter / gabor_filter.max() * 0.5 + 0.5, "Your Gabor Kernel")
```

## Exercise 1 (4 points)

The next step is to apply the filter $f(x,y)$ to a grayscale image $I(x,y)\in\mathcal{R}^{W\times H}$ using convolution. The Gabor filter response $r(x,y) = I(x,y) * f(x,y)$ has a sign that depends on the mutual orientation of the filter and the gradient. In this excercise the sign does not matter and therefore we return absolute value of the response $r_a(x,y) = |r(x,y)|$.

The filter response is expected to have the same resolution as the input image. However, we need the GIST descriptor to be relatively small. Therefore, we bin the response values into a `grid_size` $\times$ `grid_size` 2D histogram $H(i,j)$ where each bin contains the mean of $r_a(x,y)$ for the pixels that are covered by that bin.

### Task
First, implement a method that for a given grayscale image returns an absolute Gabor filter response.

Second, implement a method that bins this response into a 2D histogram as defined above. This function should work for both grayscale images and RGB color images (in which case the binned image should also have 3 color channels).

### Tips
- You can assume that the image size is divisible by the `grid_size`.
- You are allowed to use `scipy.signal.convolve2d(image, kernel, mode="same")`
- Exlore how the response of the filter changes for different parameters.

<!-- grade_id: exercise1_code, points: 0, grade: false, solution: true -->
```python
def get_absolute_response(image, filter_size, sigma, f, alpha):
    gabor_filter = build_gabor_filter(filter_size, sigma, f, alpha)
    # TODO: Compute the absolute value of the gabor filter applied to the image.
    abs_response = image.copy()
    # YOUR CODE HERE
    return abs_response, gabor_filter
    
def bin_values(image, grid_size):
    if len(image.shape) == 2:
        # Grayscale image
        binned_values = np.zeros((grid_size, grid_size))
    else:
        # Color image
        binned_values = np.zeros((grid_size, grid_size, image.shape[2]))
    
    # TODO: Bin the input image as described in the exercise description.
    # YOUR CODE HERE
    return binned_values

grayscale_image = helpers.imread_normalized_float_grayscale(os.path.join(local_data_folder, "grad_test.png"))
response, gabor_filter = get_absolute_response(grayscale_image, 32, 1.0/3.2, 8, np.radians(45))
binned = bin_values(response, 4)

helpers.show_images({
    'Input Image': grayscale_image,
    'Your Filter': gabor_filter / gabor_filter.max() * 0.5 + 0.5,
    'Your Response': response / response.max(),
    'Your Binned': plt.get_cmap("gray")(binned / max(binned.max(), 0.0000001))
}, nrows=1, ncols=4)

color_image = helpers.imread_normalized_float(os.path.join(local_data_folder, "opencountry_land493.jpg"))
response = bin_values(color_image, 4)
helpers.show_images({
    'Input Color Image': color_image,
    'Your Binned': response
}, nrows=1, ncols=4)
```

### Testing your solution of exercise 1
You can test your `bin_values(vals, grid_size)` function by creating a small and simple test image. Work out by hand what the binned image should look like and verify that your solution matches what you expect. The code below checks some basic properties such as the size of the returned image.

```python
# Add your own tests here
```

<!-- grade_id: exercise1, points: 4, grade: true, solution: false -->
```python
# DO NOT REMOVE, MODIFY, OR COPY THIS CELL

image = helpers.imread_normalized_float_grayscale(os.path.join(local_data_folder, "grad_test.png"))
response, gabor_filter = get_absolute_response(image, 32, 1.0/3.2, 8, np.radians(45))
assert(gabor_filter.shape == (32, 32))
assert(response.shape == image.shape)
binned = bin_values(response, 8)
assert(binned.shape == (8, 8))

image = helpers.imread_normalized_float(os.path.join(local_data_folder, "opencountry_land493.jpg"))
binned = bin_values(image, 16)
assert(binned.shape == (16, 16, 3))

```

## Exercise 2 (5 points)
Finally, a *GIST* descriptor is formed by concatenation of flattened histograms $H_k$ for Gabor filters $g_k$ with different scales and orientations. Here is an example of a small filter bank with 3 different sizes $s$,= and 4 orientations $\alpha$:

<img src="gabors_4x3.png" alt="Gabor filters" style="width:60%" />


If we further consider a histogram size of $2\times2$, we obtain a GIST descriptor as a 1D vector with $3\times4\times2\times2 = 48$ values (for this particular choice of parameters). 
The GIST vector features one entire histogram row after another and this is repeated first for every orientation $\alpha$ and then for every scale $s$, so that the order is:
$$
[s_0,\alpha_0,y_0,x_0], [s_0,\alpha_0,y_0,x_1],\cdots[s_0,\alpha_0,y_1,x_0],\cdots[s_0,\alpha_1,y_0,x_0],\cdots[s_1,\alpha_0,y_0,x_0],\cdots[s_S,\alpha_A,y_Y,x_X].
$$

The GIST alone only describes the distribution of gradients in the image. To make it more powerfull, we append one more histogram describing the distribution of the colors in the image similarly as shown for the collection-based inpainting in the lecture.
This color histogram is built the same way as previously but instead of the filter response it averages the RGB value of the image. When appeneded at the end of the descriptor, it adds another `grid_size` $\times$ `grid_size` $\times 3$ values that correspond to a flattened 2D color histogram with an order:

$$
[y_0,x_0,R],[y_0,x_0,G],[y_0,x_0,B],\cdots[y_0,x_1,R],[y_0,x_1,G],[y_0,x_1,B],\cdots[y_1,x_0,R],[y_1,x_0,G],[y_0,x_1,B],\cdots
$$

The final descriptor for this assignment is then defined as
$\mathbf{d} = [\textrm{GIST} \cdot w_{grad}, H_{RGB}]$
where $w_{grad}$ is additional weight scalar that scales the contribution of the gradients.

### Task
Implement a method that computes a full descriptor of given RGB image including the RGB appendix. 
Cover all combinations of sizes $s$ and orientations $\alpha$.

### Tips
- The histogram method developed in previous excercise can be easily adapted to compute the RGB histogram.
- Visualize the individual filters and their responses in a similar grid as shown in the exercise to better understand what each filter scale and orientation does.

**NOTE**: For grading we use the reference implementation of `get_absolute_response()` and `bin_values()` to ensure that you are not punished for mistakes in the previous exercises.

<!-- grade_id: cell-36f1a57a16bd2ebd, points: 0, grade: false, solution: true -->
```python
def image_rgb_to_grayscale(image):
    return cv2.cvtColor(image, cv2.COLOR_RGB2GRAY)

def build_gist_descriptor(image, filter_sizes, alphas, sigma, f, grid_size, grad_weight):
    # Use image_rgb_to_grayscale() to convert the RGB image to grayscale.
    descriptor = np.zeros((len(filter_sizes) * len(alphas) + 3) * grid_size*grid_size)
    image_grayscale = image_rgb_to_grayscale(image)
    # TODO: fill the GIST descriptor as descriptor as described in the exercise description.
    # YOUR CODE HERE
    return descriptor
```

### Testing your solution of exercise 2
Below is a visualization of the different parts of the image descriptor that you've generated.

```python
# Add your own tests here.
```

```python
image = helpers.imread_normalized_float(os.path.join(local_data_folder, "spatenv/opencountry_land228.png"))

filter_sizes = np.array([9,17,33,65])
alphas = np.radians(np.arange(6, dtype=float)/6*180)
sigma = 1.0 / 3.2
f = 8
grid_size = 4
grad_weight = 20
descriptor = build_gist_descriptor(image, filter_sizes, alphas, sigma, f, grid_size, grad_weight)

# Visualize.
dsize = grid_size**2
s_maxs = np.array([descriptor[i*(dsize*len(alphas)):(i+1)*(dsize*len(alphas))].max() for i in range(len(filter_sizes))])
panels = {}
for i, (s, a) in enumerate(itertools.product(filter_sizes, alphas)):
    d = descriptor[i*dsize:(i+1)*dsize].reshape(grid_size, grid_size)
    panels[f's={s} a={np.degrees(a):.0f}'] = plt.get_cmap('viridis')(d / s_maxs[filter_sizes==s])
panels['RGB'] = descriptor[-(dsize*3):].reshape(grid_size,grid_size,3)
skeys = ['s=9 a=0','s=17 a=60', 's=33 a=90', 's=65 a=120', 'RGB']
panels = {f'Your {k}': panels[k] for k in skeys}
panels.update({ "Input": image })
helpers.show_images(panels, ncols=len(panels), col_height=2.2)
```

<!-- grade_id: exercise3, points: 5, grade: true, solution: false -->
```python
# DO NOT REMOVE, MODIFY, OR COPY THIS CELL

```

## Exercise 3 (4 points)

We can now use the resulting descriptor to quickly find a set of $k$ most similar images in potentially large image collection which would take a long time to compare pixel-by-pixel. 

In this exercise we will find the most similar image by minimizing the sum of squared distances between the query image descriptor $\mathbf{d}$ and the collection image descriptors $\mathbf{d}_i$ such that the index of the most similar image $j$ is
$$
j = \arg\min_{i} ||\mathbf{d} - \mathbf{d}_i||^2_2.
$$

This is useful for exploring large image collections where a search engine can suggest other images similar to the one we are looking at. It can also be used for a coarse categorization of images (coast, mountains, road, city,...) and in the lecture we have seen that a similar image can be used for image inpainting. 

Note, that for an effective implementation, we must ignore the pixels that are marked as invalid in a provided inpainting mask $M$. That means that every convolution has to carefully avoid sampling the invalid pixels which results in a more complex and slow implementation This is why in this exercise we simulate similar effect by modifying the query image.

### Task

Implement a method that replaces marked pixels in a given RGB image $\mathbf{I} \in \mathcal{R}^{H \times W \times 3}$ with pixels from the most similar image in a provided image collection. 
The marked pixels are encoded by zeros in the provided mask $M \in [0,1]^{H \times W}$.
Find the most similar image by minimizing the L2 descriptor norm as shown above.
To compute the query image descriptor, use a modified query image generated from the input image by replacing the invalid pixels ($M = 0$) by a uniform gray color ($[0.5, 0.5, 0.5]$).
Return the modified query image, the most similar collection image and the inpainted image.


### Tips
- The descriptors for all images in the dataset are provided and do not need to be recomputed.
- Visualize several of the most similar and also several of the least similar images. 
- Test the query with your own image. Note that the metric is only effective if the collection contains a very close match of the query image. In practice, that means that collections of milions of images are needed. In this example we only use a very small collection of pre-selected image samples.

Images courtesy of [8 Scene Categories Dataset](http://people.csail.mit.edu/torralba/code/spatialenvelope/)

**NOTE**: For grading we use the reference implementation of `build_gist_descriptor()` to ensure that you are not punished for mistakes in the previous exercises.

<!-- grade_id: exercise4_code, points: 0, grade: false, solution: true -->
```python
def inpaint_image(image, mask, dataset_descriptors, dataset_files, 
                 filter_sizes, alphas, sigma, f, grid_sizes, grad_weight):
    # Use helpers.imread_normalized_float(dataset_files[i]) to read an image i
    masked_query_image = image.copy()
    nearest_image = helpers.imread_normalized_float(dataset_files[0]) # First image in the dataset
    inpainted_image = image.copy()

    # TODO 1: Set the masked pixels in masked_query_image to gray ([0.5, 0.5, 0.5])
    # TODO 2: Compute the GIST descriptor for masked_query_image and select the closest descriptor from dataset_descriptors.
    # TODO 3: Load the corresponding image from dataset_files (use helpers.imread_normalized_float) into nearest_image
    # TODO 4: Use nearest_image to inpaint the masked part of the image (store the result in inpainted_image).
    # YOUR CODE HERE
    return masked_query_image, nearest_image, inpainted_image

filter_sizes = np.array([9,17,33,65])
alphas = np.radians(np.arange(6, dtype=float)/6*180)
sigma = 1.0 / 3.2
f = 8
grid_size = 4
grad_weight = 20

dataset_files = sorted(helpers.list_files_in_folder(os.path.join(local_data_folder, "spatenv"), filter="*.png"))
dataset_descriptors = np.load(os.path.join(local_data_folder, 'spatenv_descs.npy'))
image = helpers.imread_normalized_float(os.path.join(local_data_folder, "opencountry_land493.jpg"))
mask = np.zeros(image.shape[:2])
mask[:image.shape[0]//2] = 1

query, nearest, inpainted = inpaint_image(image.copy(), mask, dataset_descriptors, dataset_files,
                                        filter_sizes, alphas, sigma, f, grid_size, grad_weight)
                                        
helpers.show_images({
    'Your Input': image,
    'Your Query': query,
    'Your Nearest': nearest,
    'Your Result': inpainted
}, nrows=1, ncols=4, col_height=3)
```

### Testing your solution of exercise 3
Have a look at the set of images in the data set (`datasets/week3/collections/spatenv/` folder). Is the image picked by your algorithm close to the query image (compared to the other images in the data set)? Try different query images and check if this is still the case.

The code below checks some basic properties such as the size of the returned images.

```python
# Add your own tests here
```

<!-- grade_id: exercise4, points: 4, grade: true, solution: false -->
```python
# DO NOT REMOVE, MODIFY, OR COPY THIS CELL

filter_sizes = np.array([9,17,33,65])
alphas = np.radians(np.arange(6, dtype=float)/6*180)
sigma = 1.0 / 3.2
f = 8
grid_size = 4
grad_weight = 20

dataset_files = sorted(helpers.list_files_in_folder(os.path.join(local_data_folder, "spatenv"), filter="*.png"))
dataset_descriptors = np.load(os.path.join(local_data_folder, 'spatenv_descs.npy'))
image = helpers.imread_normalized_float(os.path.join(local_data_folder, "opencountry_land493.jpg"))
mask = np.zeros(image.shape[:2])
mask[:image.shape[0]//2] = 1

import warnings
with warnings.catch_warnings():
    warnings.simplefilter('ignore')
    query, nearest, inpainted = inpaint_image(image.copy(), mask, dataset_descriptors, dataset_files,
                                        filter_sizes, alphas, sigma, f, grid_size, grad_weight)
assert(query.shape == image.shape)
assert(nearest.shape == image.shape)
assert(inpainted.shape == image.shape)

```